---
title: Learning Materials
---

## Notebooks

Tutorial materials can be launched from the course homepage:

::::{attention} Course homepage for registered students

<https://canvas.cityu.edu.hk/courses/67502>

::::

If you have not yet registered the course, visit the static site:

::::{attention} Static Jupyter book

<https://ccha23.github.io/cs5483i25b/>

::::

:::::{table} An example of a notebook link on the course homepage.
:label: tbl:notebook-link

<table>
<div style="border: 1px solid black; padding: 10px;">
<p title="Title: Tutorial">Tutorial notebooks:</p>
<ul>
  <li>
    <a title="Subtitle: Notebook link" href="https://www.cs.cityu.edu.hk/~ccha23/cs5483/nb/?path=Tutorial1/Learning_Materials.ipynb" target="_blank">Learning Materials</a>
  </li>
  <li>
    <a title="Subtitle: Notebook link" href="https://www.cs.cityu.edu.hk/~ccha23/cs5483/nb/?path=Tutorial1/DIVE_AI.ipynb" target="_blank">DIVE AI</a>
  </li>
  <li>
    <a title="Subtitle: Notebook link" href="https://www.cs.cityu.edu.hk/~ccha23/cs5483/nb/?path=Tutorial1/Data_Files.ipynb" target="_blank">Data Files</a>
  </li>
</ul>
</div>
</table>

:::::

## Notebooks

Tutorials are written in the form of [Jupyter notebooks](https://jupyter.org/), which provide an interactive [literate programming](https://en.wikipedia.org/wiki/literate_programming) experience:

- Jupyter notebooks allow programs to be mixed with other contents, including figures, videos, and formatted textual explanations.
- The notebooks can also be opened in a Jupyter server, which enables users to edit and run the programs in addition to other contents.

### How to run notebook cells?

If you already have your notebook opened in a Jupyter server, let's run some code. Otherwise, jump to the [next section](#sec:fetch) to learn how to fetch and open notebooks.

To run a cell, select the cell and press <kbd>Shift + Enter</kbd>. The cell prompt with change from `[ ]` to `[*]`. When the run completes, the asterisk will be changed to a number like `[1]`. Try running the following cell to import the [Mathematical Animation Engine (Manim)](https://docs.manim.community/).[^import] 

[^import]: Importing a library in programming is like bringing a toolbox into your workshop, giving you access to all the tools you need to complete your tasks.

```python
import os

if not os.getenv(
    "NBGRADER_EXECUTION"
):  # Skip the code or auto-grading may take too long to complete.
    import manim
```

After the import is complete,  run the following cell to create an animation:

```python
%%manim -qm --progress_bar=none --disable_caching --flush_cache -v ERROR Welcome
class Welcome(manim.Scene):
    def construct(self):
        self.play(manim.Write(manim.Text("Welcome to CS5483!")))
```

As you can see, the code creates a scene that displays the message `Welcome to CS5483!` with a writing animation.[^manim-magic] To verify that the program runs live, feel free to modify the message to anything you want.

[^manim-magic]: If you are interested in what the first line does, it sets the video quality to medium while disabling caching and progress bars, and setting the verbosity to reporting only ERROR but not other critical information.

(sec:fetch)=
### How to fetch and open notebooks?

To fetch the notebooks, follow the links to the notebooks on the course homepage such as the one in [](#tbl:notebook-link).

::::{seealso} What does the notebook link do?

The notebooks reside in a [GitHub repository](https://github.com/dive4dec/cs5483_24b/blob/main/README.md). Accessing a notebook link will

1. use [LTI Authenticator](https://ltiauthenticator.readthedocs.io) to launch the JupyterHub as an LTI external tool from the course homepage, and
2. use [NbGitPuller](https://jupyterhub.github.io/nbgitpuller/index.html) to [git-pull](https://git-scm.com/docs/git-pull)s all the files in the repository, merge them with any existing files stored under the course folder `cs5483_24b` in your home directory without overwriting your changes, and open the notebook path specified in the JupyterLab interface.

Manually fetching the notebooks through other means may not create the local repository required for pulling new notebooks or changes to existing notebooks. To ensure that you have the latest versions of the notebooks with proper version control, you should use the provided notebook links.

::::

After clicking a notebook link, you may be asked to 

- login with your CityU account if you are not yet logged into Canvas, and
- specify a server option if you do not yet have a running Jupyter server.

Select the `Default` option from the list of available Jupyter servers, and click <kbd>Start</kbd> to begin your session.

::::{seealso} What is a Jupyter server?
:class: dropdown

The course's programming environment is conveniently accessible via a [JupyterHub][jh] server, enabling remote access without requiring special software installations, thanks to [Project Jupyter](https://en.wikipedia.org/wiki/Project_Jupyter). This allows you to write and run programs on mobile devices using just a web browser. Each server runs the [Ubuntu](https://en.wikipedia.org/wiki/Ubuntu) operating system in a [container](https://en.wikipedia.org/wiki/Containerization_(computing)). See the [Dockerfile](https://github.com/dive4dec/jupyter/blob/main/cs5483nb/Dockerfile) for the setup.

::::

[jh]: https://jupyterhub.readthedocs.io/en/stable/

Each student and project group will have their own [Jupyter server](https://jupyter-server.readthedocs.io/en/latest/) with individual computing resources including CPU, memory, and GPU:

- The GPU server option(s) are available initially for testing. Once the project groups are formed, the GPU server option(s) will be available to project groups to work on projects that rely on GPU resources.
- You can switch between different server options by restarting your server and selecting the desired option. Restarting the server is simple and can be done by
  1. selecting `Hub Control Panel` under the `File` menu, and
  2. click `Stop My Server` and then `Start My Server`.

:::{tip} Troubleshooting

In case the server failed to spawn, refresh the page or restart your browser for the javascript to load completely in your browser. 

- If you failed to spawn a GPU server, it is likely due to the **GPU being fully utilized**. You can click the `Home` link to start the `Default` server.
- If you are unable to spawn the Default server, you may have **exceeded your storage limit**. If this occurs, please contact us so we can assist with clearing up storage space to resolve the issue.

:::

If the server is spawned successfully, the [JupyterLab](https://jupyterlab.readthedocs.io) interface should appear. For instance, to see the available storage and disk usage, you can run the [commands](https://www.gnu.org/software/coreutils/manual/html_node/Disk-usage.html) within a notebook cell as follows:

```python
if not os.getenv("NBGRADER_EXECUTION"):
    # `!` is used to execute a shell command from within the notebook
    !df .
    !du -ahd1 ${HOME} | sort -rh | head -n10
```

Monitor your disk usage to avoid exceeding the storage limit. The tutorial notebooks are under the course directory:

```python
if not os.getenv("NBGRADER_EXECUTION"):
    # Construct the course directory path from the two environment variables
    course_dir = os.path.join(os.getenv("HOME"), os.getenv("COURSE_ID"))
    print(f"Course directory: {course_dir}")
    !ls "{course_dir}"
```

To prevent data loss, it's essential to regularly download or backup your work. You may use the [JupyterLab Archiver extension](https://github.com/jupyterlab-contrib/jupyter-archive) to download an entire folder as a zip file.

:::{tip} Troubleshooting issues with JupyterLab

You may occasionally run into issues. There are thousands of open issues reported to the JupyterLab repository, where you may find related issues and a kludge (temporary workaround):

- <https://github.com/jupyterlab/jupyterlab/issues>

- If the JupyterLab interface fails to show, there may be issue loading or running the necessary javascripts. Try to ensure a good network connection, refresh the page, or restart your browser for the javascript to load properly in your browser. You may also inspect the javascript error in your browser's developer tools for details.

- If you are unable to open the notebook or you believe your notebooks are corrupted:  

  1. Move the existing course folder, if any, to a different location.
  2. Click the git-pull links from the course homepage to re-pull the notebooks.
  3. Copy your solutions from your original to the new notebooks.


:::

::::{tip} Access via bookmark
:class: dropdown

If you access JupyterHub from a [bookmarked link](https://dive.cs.cityu.edu.hk/cs5483_24a/), you may see a Login box as shown in [](#fig:login-box). Since the bookmarked link does not perform LTI login through Canvas, you will need to log in separately:

  1. Enter your [EID](https://www.cityu.edu.hk/esu/eid.htm) and password in the fields `Username` and `Password` respectively.
  1. Click the `Sign in` button.

Note that the bookmarked link also does not pull updates to notebooks from the remote Git repository.

:::{figure} images/login.dio.svg
:name: fig:login-box
:alt: Login box
:align: left

Login box
:::

::::

## Problem Sets

The tutorial notebooks consisting of problems that provide hands-on practice to prepare you for the individual **Project 1** and group **Project 2**.

Each problem set is *due the night before the following lecture*, i.e., you normally have a week to work on the problem set.

- You can check the due dates of all the labs from the course homepage. 
- You may seek help from us or your classmates. However, you must write your own solution and indicate who your collaborators are by including their fullnames or EIDs such as
  ```
  COLLABORATORS: xfwong2, mklee3, GPT-4o
  ```
  The policy applies to GenAI as well.

### How to complete an assignment?

The notebooks can be edited in JupyterLab to include your answers for submission. If this is your first time using JupyterLab, take a look at the [official video tutorial](https://youtu.be/A5YyoCKxEOU?si=cdaLdWlF35Gwf0KH):


::::{seealso} How to use JupyterLab?
:class: dropdown

:::{iframe} https://www.youtube.com/embed/A5YyoCKxEOU
:width: 100%
:align: left

Official video tutorial on Jupyter
:::
::::

For more advanced features:

- Checkout the `Help` menu items
    - `Show Keyboard Shortcuts` to try some of the shortcuts to see their effect, and
    -  `Jupyter Reference` to open the [user guide](https://jupyterlab.readthedocs.io/en/latest/user/interface.html) in a new tab.
- Try also the [MyST Markdown](https://mystmd.org/guide/typography) syntax to format your notes.

::::{seealso} An alternative interface: VSCode

You may also use a highly customizable editor Visual Studio Code (VS Code) to open a notebook:

- Click `File → New Launcher → VS Code`.
- Click the menu icon on the left and select `Open Folder` and select `cs1302_24a`.

In the file explorer, you can navigate to a notebook to open it. However, to properly run the notebook, you would also need to start the Kernel by choosing an appropriate Python environment such as `base`, which is the default [(conda) environment](https://conda.io/projects/conda/en/latest/dev-guide/api/conda/base/index.html) for our JupyterLab setup. Unlike JupyterLab, you may install additional extensions yourself to enrich the interface.[^vscode]

::::

[^vscode]: The VSCode interface is actually [code-server](https://github.com/coder/code-server), so some VSCode extensions may not be listed under the extension panel. You may still try to install them using the command `install-vscode-extension` in a terminal.

As a example, you will write a ["Hello, World!"](https://en.wikipedia.org/wiki/%22Hello,_World!%22_program) program in python, which says Hello to the world:

:::::{exercise} Hello World
:label: ex:Hello-World

Complete the program to print the message `Hello, World!`. 

::::{hint}
:class: dropdown

One possible solution is:

:::{code-block} python
:name: code:Hello-World
:caption: Python, "Hello, World!"

def say_hello():
    print("Hello, World!")
    
say_hello()
:::

:::{caution}
You must use a tab or 4 spaces to indent the second line of code `print(...)`.
:::

::::

:::::

The following code cell is a [solution cell](https://nbgrader.readthedocs.io/en/stable/configuration/student_version.html#default-behavior). You can simply expand the hint above and copy the answer to the solution cell instead.

<!-- grade_id: helloworld, points: 0, grade: false, solution: true -->
```python
def say_hello():
    # YOUR CODE HERE
    raise NotImplementedError

say_hello()
```

:::{attention}

It's important to follow certain guidelines when writing your answers or code in the notebooks:

- Only provide your answers in the cells that are designated for this purpose, which are usually labeled with `YOUR CODE HERE` or `YOUR ANSWER HERE`.
- For coding exercises, be sure to remove the line `raise NotImplementedError()` from the cell before submitting your work.
- Do not clone or duplicate the answer cells, as this can cause issues with version control and grading.

:::

To test your program:

- Run your program by selecting the solution cell and press <kbd>Shift + Enter</kbd>.
- Then, run the following visible test to check whether your program prints the correct message.

<!-- grade_id: test-helloworld, points: 1, grade: true, solution: false -->
```python
# Run this test cell right after running your "Hello, World!" program.
import io
import sys

old_stdout, sys.stdout = sys.stdout, io.StringIO()
say_hello()
printed = sys.stdout.getvalue()
sys.stdout = old_stdout
assert printed == "Hello, World!\n"
```

The test returns an assertion error if your program does not print the correct message.

::::{tip}

- You can repeatedly modify your solution and run the test cell until your solution passes the test. There is no mark penalty in failing the test before submission. 
- You should try to understand the error messages from a failed test.
- To assess your solution thoroughly, we will run new tests hidden from you after you have submitted your notebook. Therefore, *you should ensure your solution works in general rather than just the visible tests*.
- If you open the same notebook multiple times in different browser windows, be careful in making changes in different windows as you may overwrite your own changes.
- If your notebook fails to run any code, the kernel might have died. You can restart the kernel with `Kernel → Restart`. If restarting fails, check your code cells to see if there is anything that breaks the kernel such as an infinite loop.

::::

### How to submit?

:::{important} Grading policies

- Late submissions will not be accepted without valid justifications.
- You can submit your assignment repeatedly before the deadline without penalty, so it is recommended to submit your assignment in advance. Please note that you are responsible for any issues related to failed submissions due to network outages that occur too close to the deadline.

:::

Before you submit, it is a good idea to make sure everything runs as expected:

1. **Git-pull the notebooks**: Follow any one of the link on the course homepage to a notebook, which will git-pull any updates/corrections to (all) your notebooks.
1. **Save the notebook**: Unsaved changes are not submitted, even though they are in the memory.
1. **Restart the kernel**: `Kernel`$\to$ `Restart Kernel...` to have a clean state before running visible tests.
1. **run all cells**: `Run`$\to$ `Run All Cells` to double check the results of the visible tests are as expected.

Re-executing cells in order in a Jupyter Notebook ensures that all dependencies and state changes are correctly applied, preventing misleading results from out-of-order execution.

To submit your notebooks:

1. Select the menu item `Nbgrader → Assignment List`.
1. Expand the Lab folder and click the `validate` button next to the notebook(s) to check if all the visible tests pass.
1. Click `Submit` to submit your notebook.

::::{tip}

1. By default, JupyterLab autosaves an opened notebook every 120 seconds. This configuration can be changed at `Settings`$\to$`Settings Editor`$\to$`Document Manager`$\to$`Autosave Interval`.
2. An opened file with unsaved changes has a solid circle next to its tab label. You can manually save the file using the save button in the toolbar. It is also important to manually save your notebook before submission because changes not saved to disk will not be submitted.

3. Validation is important as it runs the notebook cells in order starting from a clean state. Re-executing cells in order in a Jupyter Notebook ensures that all dependencies and state changes are correctly applied, preventing misleading results from out-of-order execution. The autograding behavior may also be modified by the [environment variable `NBGRADER_EXECUTION`](https://nbgrader.readthedocs.io/en/latest/user_guide/faq.html#can-i-modify-the-notebook-behavior-during-autograding-or-validation).

::::

::::{seealso} What is Nbgrader?
:class: dropdown

[Nbgrader](https://nbgrader.readthedocs.io/en/stable/) is a package for grading notebook assignments. It allows students to submit their notebooks directly through the JupyterHub server. Submitted notebooks can be both auto-graded with pre-defined test cases and manually graded with custom feedback. After grading is complete, you can check your scores and access the feedback using the same interface.

::::

::::{caution}

A common and costly mistake is forgetting to validate/submit the notebooks after completing the problem set. The grading may also fail under the following circumstances:

- The notebook files have been renamed, for example, `Setup.ipynb` being changed or copied to `Setup (Copy).ipynb`.
- An HTML file exists with the same name as a notebook, for example, `Setup.html`.
- The file size is too large, e.g., exceeds `100MB`.
- The code takes too long to run or requires an excessive amount of memory.

It is essential to ensure that your submission meets these guidelines to avoid any issues with grading.

::::